---
title: "Manipulating data with `dplyr`"
author: "Julian Faraway and Benjamin A. Robinson"
date: "`r format(Sys.time(), '%d %B %Y')`"
output: learnr::tutorial
runtime: shiny_prerendered
---

```{r setup, include = FALSE}
library(learnr)
library(gradethis)
tutorial_options(exercise.checker = gradethis::grade_learnr)
library(R.utils)
setOption("gradethis_glue_correct", "{ .message} {.correct}")
setOption("gradethis_glue_incorrect", "{ .message} { .incorrect}")
library(tidyverse)
load("hospds.rda")
```

## Setup

The `dplyr` package is part of the `tidyverse`. It is useful for manipulating data. See <https://www.tidyverse.org> for more details.

We need to start by loading the `tidyverse` and the data.

```{r}
library(tidyverse)
load("../../../data/hospds.rda")
```

If you continue immediately from the previous exercise, the data and package would already be loaded. But is sensible to make your scripts work individually.

## Filter

The `dplyr` package uses "verbs" to manipulate data. The first verb we learn is `filter` which chooses rows of the data based on a condition. Here is an example:

```{r}
filter(hospds, City == "Ulaanbaatar")
```

We find all cases from Ulaanbaatar. We see there are 17,482 cases.

It is common to use a pipe in dplyr computations. The '%>%' operator is used:

```{r}
hospds %>% filter(City == "Ulaanbaatar")
```

The output from the object on the left of the pipe (`hospds`) is used as input to the expression on the right of the pipe. We can chain these together:

```{r}
hospds %>% filter(City == "Ulaanbaatar") %>% filter(Age == 90)
```

These are 5 90 year olds admitted in Ulaanbaatar. We can do this more neatly as:

```{r}
hospds %>% filter(City == "Ulaanbaatar" & Age == 90)
```

The `slice` verb is similar to `filter`. It selects the rows using the row number. The `:` operator generates sequences of numbers. For example:

```{r}
10:20
```

We can select rows 10 to 20 in the data using

```{r}
hospds %>% slice(10:20)
```

### Exercises

1. Find all the patients admitted to Hospital no. 613.

```{r filter_ex_1, exercise = TRUE}

```

```{r filter_ex_1-hint-1}
hospds %>% filter(___)
```

```{r filter_ex_1-hint-2}
hospds %>% filter(___ == ___)
```

```{r filter_ex_1-hint-3}
hospds %>% filter(Hospital == 613)
```

```{r filter_ex_1-solution}
hospds %>% filter(Hospital == 613)
```

```{r filter_ex_1-check}
grade_code(
  correct = "Correct!",
  incorrect = "Try again"
)
```
How many are there? (Write some code that outputs the number)

```{r filter_ex_1b, exercise = TRUE}

```

```{r filter_ex_1b-solution}
hospds %>% filter(Hospital == 613) %>% count()
```
```{r filter_ex_1b-check}
grade_result(
  pass_if(~ .result == hospds %>% filter(Hospital == 613) %>% count(), "Correct!"),
  default_message = "Try again"
)
```

2. Find all the patients admitted to Hospital no. 2162636 who are also older than 21.

```{r filter_ex_2, exercise = TRUE}

```

```{r filter_ex_2-hint-1}
hospds %>% filter(___ & ___)
``` 

```{r filter_ex_2-hint-2}
hospds %>% filter(Hospital ___ & Age ___)
```

```{r filter_ex_2-hint-3}
hospds %>% filter(Hospital == 2162636 & Age > 21)
```

```{r filter_ex_2-solution}
hospds %>% filter(Hospital == 2162636 & Age > 21)
```

```{r filter_ex_2-check}
grade_result(
  pass_if(~ identical(.result, hospds %>% filter(Hospital == 2162636 & Age > 21)) , "Correct!"),
  fail_if(~ identical(.result, hospds %>% filter(Hospital == 2162636)) , "You also need to filter by age"),
   fail_if(~ identical(.result, hospds %>% filter(Age > 21)) , "You also need to filter by hospital"),
  default_message = "Try again"
)
```

3. Find the first three patients who are aged 79 years. You will need to use both filter and slice.

```{r filter_ex_3, exercise = TRUE}

```

```{r filter_ex_3-hint-1}
hospds %>% filter(___) %>% slice(___)
```

```{r filter_ex_3-hint-2}
hospds %>% filter(Age == 79) %>% slice(___)
```

```{r filter_ex_3-hint-3}
hospds %>% filter(Age == 79) %>% slice(1:3)
```

```{r filter_ex_3-solution}
hospds %>% filter(Age == 79) %>% slice(1:3)
```

```{r filter_ex_3-check}
grade_result(
  pass_if(~ identical(.result, hospds %>% filter(Age == 79) %>% slice(1:3)), "Correct!"),
  fail_if(~ identical(.result, hospds %>% filter(Age == 79)) , "You also need to use slice to get the first 3 patients"),
   fail_if(~ identical(.result, hospds %>% slice(1:3)) , "You also need to filter by age"),
  default_message = "Try again"
)
```


## Select

The `select` verb is used for selecting variables(columns) in the data. 

For example, we can choose two columns:

```{r}
hospds %>% select(Hospital, City)
```

or everything except these two columns:

```{r}
hospds %>% select(-Hospital, -City)
```

We can use a modifier if we only want variables of a particular type For example, suppose we want only the numeric type variables:

```{r}
hospds %>% select_if(is.numeric)
```

Sometimes we want a range of variables. We can do this by using the range operator `:`.

```{r}
hospds %>% select(Gender:Day)
```

More *select helpers* can be found at <https://dplyr.tidyverse.org/reference/select_helpers.html>


### Exercises

1. Create a dataset with only the `Gender` and `City` variables.

```{r select_ex_1, exercise = TRUE}

```

```{r select_ex_1-hint-1}
hospds %>% select(___, ___)
```

```{r select_ex_1-hint-2}
hospds %>% select(Gender, City)
```

```{r select_ex_1-solution}
hospds %>% select(Gender, City)
```

```{r select_ex_1-check}
grade_result(
  pass_if(~ identical(.result, hospds %>% select(Gender, City)), "Correct!"),
  fail_if(~ identical(.result, hospds %>% select(Gender)) , "You also need to include the City variable"),
   fail_if(~ identical(.result, hospds %>% select(City)) , "You also need to include the Gender variable"),
  default_message = "Try again"
)
```

2. Create a dataset with only character variables.

```{r select_ex_2, exercise = TRUE}

```

```{r select_ex_2-hint-1}
hospds %>% ___(is.character)
```

```{r select_ex_2-hint-2}
hospds %>% select_if(is.character)
```

```{r select_ex_2-solution}
hospds %>% select_if(is.character)
```

```{r select_ex_2-check}
grade_code(
  correct = "Correct!",
  incorrect = "Try again"
)
```

3. Create a dataset with only the cities of those aged 50.

```{r select_ex_3, exercise = TRUE}

```

```{r select_ex_3-hint-1}
hospds %>% filter(___) %>% select(___)
```

```{r select_ex_3-hint-2}
hospds %>% filter(Age == 50) %>% select(___)
```

```{r select_ex_3-hint-3}
hospds %>% filter(Age == 50) %>% select(City)
```

```{r select_ex_3-solution}
hospds %>% filter(Age == 50) %>% select(City)
```

```{r select_ex_3-check}
grade_result(
  pass_if(~ identical(.result, hospds %>% filter(Age == 50) %>% select(City)), "Correct!"),
  fail_if(~ identical(.result, hospds %>% filter(Age == 50)) , "You also need to use `select` to get a dataset with only the cities"),
   fail_if(~ identical(.result, hospds %>% filter(Age == 50) %>% select(City, Age)) , "You have created a dataset with both cities and ages. You only need to keep the City variable."),
  fail_if(~ identical(.result, hospds %>% select(City, Age)) , "You need to filter by age and then keep only the City variable."),
  fail_if(~ identical(.result, hospds %>% select(City)) , "You need to filter by age before selecting the city variable."),
  default_message = "Try again"
)
```


## Arrange

The `arrange` verb sorts the data. Here is an example.

```{r}
hospds %>% arrange(Age)
```

The data are now sorted with those aged zero first. We can reverse the order:

```{r}
hospds %>% arrange(desc(Age))
```

### Exercises

1. Sort the data by hospital number, with the highest hospital number first.

```{r arrange_ex_1, exercise = TRUE}

```

```{r arrange_ex_1-hint-1}
hospds %>% arrange(___)
```

```{r arrange_ex_1-hint-2}
hospds %>% arrange(___(Hospital))
```

```{r arrange_ex_1-hint-3}
hospds %>% arrange(desc(Hospital))
```

```{r arrange_ex_1-solution}
hospds %>% arrange(desc(Hospital))
```

```{r arrange_ex_1-check}
grade_result(
  pass_if(~ identical(.result, hospds %>% arrange(desc(Hospital))), "Correct!"),
  fail_if(~ identical(.result, arrange(Hospital)) , "You have sorted by hospital number with the smallest number first."),
  default_message = "Try again"
)
```

What is the highest hospital number? You can use a combination of `arrange`, `select`, and `slice` to find this.

```{r arrange_ex_1b, exercise = TRUE}

```

```{r arrange_ex_1b-hint-1}
hospds %>% arrange(___) %>% select(___) %>% slice(___)
```

```{r arrange_ex_1b-hint-2}
hospds %>% arrange(desc(Hospital)) %>% select(___) %>% slice(___)
```

```{r arrange_ex_1b-hint-3}
hospds %>% arrange(desc(Hospital)) %>% select(Hospital) %>% slice(___)
```

```{r arrange_ex_1b-hint-4}
hospds %>% arrange(desc(Hospital)) %>% select(Hospital) %>% slice(1)
```

```{r arrange_ex_1b-solution}
hospds %>% arrange(desc(Hospital)) %>% select(Hospital) %>% slice(1)
```

```{r arrange_ex_1b-check}
grade_result(
  pass_if(~ identical(.result, hospds %>% arrange(desc(Hospital)) %>% select(Hospital) %>% slice(1)), "Correct!"),
  fail_if(~ identical(.result, hospds %>% arrange(Hospital) %>% select(Hospital) %>% slice(1)) , "You have found the lowest hospital number."),
  fail_if(~ identical(.result, hospds %>% arrange(desc(Hospital)) %>% slice(1)) , "Now use select to get only the hospital number."),
  default_message = "Try again"
)
```

2. Use `filter` and `arrange` to create a data set with the only the patients in Khovd, sorted by age from oldest to youngest.

```{r arrange_ex_2, exercise = TRUE}

```

```{r arrange_ex_2-hint-1}
hospds %>% filter(___) %>% arrange(___)
```

```{r arrange_ex_2-hint-2}
hospds %>% filter(___) %>% arrange(desc(Age))
```

```{r arrange_ex_2-hint-3}
hospds %>% filter(City == "Khovd") %>% arrange(desc(Age))
```

```{r arrange_ex_2-solution}
hospds %>% filter(City == "Khovd") %>% arrange(desc(Age))
```

```{r arrange_ex_2-check}
grade_result(
  pass_if(~ identical(.result, hospds %>% filter(City == "Khovd") %>% arrange(desc(Age))), "Correct!"),
  fail_if(~ identical(.result, hospds %>% arrange(desc(Age))) , "You also need to filter by city."),
  fail_if(~ identical(.result, hospds %>% filter(City == "Khovd")) , "You also need to sort by age."),
  fail_if(~ identical(.result, hospds %>% filter(City == "Khovd") %>% arrange(Age)) , "You have sorted by age from youngest to oldest, rather than oldest to youngest."),
  default_message = "Try again"
)
```

What was the age of oldest person admitted in Khovd? You can use your previous code with `select` and `slice` to output this number.

```{r arrange_ex_2b, exercise = TRUE}

```

```{r arrange_ex_2b-hint-1}
hospds %>% filter(City == "Khovd") %>% arrange(desc(Age)) %>% select(___) %>% slice(___)
```

```{r arrange_ex_2b-hint-2}
hospds %>% filter(City == "Khovd") %>% arrange(desc(Age)) %>% select(Age) %>% slice(___)
```

```{r arrange_ex_2b-hint-3}
hospds %>% filter(City == "Khovd") %>% arrange(desc(Age)) %>% select(Age) %>% slice(1)
```

```{r arrange_ex_2b-solution}
hospds %>% filter(City == "Khovd") %>% arrange(desc(Age)) %>% select(Age) %>% slice(1)
```

```{r arrange_ex_2b-check}
grade_result(
  pass_if(~ identical(.result, hospds %>% filter(City == "Khovd") %>% arrange(desc(Age)) %>% select(Age) %>% slice(1)), "Correct!"),
  fail_if(~ identical(.result, hospds %>% arrange(desc(Age)) %>% select(Age) %>% slice(1)) , "You have found the age of the oldest person admitted in any city. You need to use filter to get the olderst person amitted in Khovd."),
  fail_if(~ identical(.result, hospds %>% filter(City == "Khovd") %>% arrange(desc(Age)) %>% slice(1)) , "You need to use select to get only the age."),
  fail_if(~ identical(.result, hospds %>% filter(City == "Khovd") %>% arrange(Age) %>% select(Age) %>% slice(1)) , "You have found the youngest person admitted in Khovd, not the oldest."),
  default_message = "Try again"
)
```

## Mutate

The `mutate` verb is used add columns to the data. Suppose we want the exact age of a patient. The `Age` is given in years but the `Month` and the `Day` should be combined with this:


```{r}
hospds %>% mutate(trueage = Age + Month/12 + Day/365) %>% select(Age, Month, Day, trueage)
```

I have selected only these variables for display to make it easier to see what has be done. We can see that the true age has been added to the data in the final column. But now try this:

```{r}
colnames(hospds)
```

We can see that the true age variable is not included. We did not save our computation. We can achieve this with:


```{r}
hospds %>% mutate(trueage = Age + Month/12 + Day/365) -> hospds
colnames(hospds)
```

We can see it at the end of the list of variable names.

### Exercises

Create a new variable which is the departure date minus the admission date. Compare it to the `TreatDays` variable.

## Summarize

The `summarise` verb computes summaries of variables. Suppose we want the mean age:

```{r}
hospds %>% summarise(mean(Age))
```

Now suppose we want the median stay (`TreatDays`):

```{r}
hospds %>% summarize(median(TreatDays))
```

As it happens, we can do the same thing with a base R command:

```{r}
median(hospds$TreatDays)
```

*Base R* refers to the commands available without loading any packages. It is powerful and can do anything that *dplyr* can do. But its syntax is more complex and less transparent.

The `summarise` command is more useful when we want more than one summary:

```{r}
hospds %>% summarise(min = min(Age), mean=mean(Age), max=max(Age))
```

Notice how we have named the three elements so they are labelled in the output.

The `summarize_at` command allows us apply several summaries to several variables. Here we compute the min, mean and max of two variables. 

```{r}
hospds %>% summarize_at(vars(Age,TreatDays),list(min=min,mean=mean,max=max))
```

The `summarize_if` command allows us to choose variables based on some characteristic. In this example, we check for numeric variables and then summarise them:

```{r}
hospds %>% summarize_if(is.numeric, list(min=min,mean=mean,max=max))
```

Notice that it would not make sense to ask for the mean of a character variable.

The `summarize_all` command applies the function to all the variables. Here is an example:

```{r}
hospds %>% summarize_all(last)
```

This asks for the last value of every variable. This function is valid for all variables so it is acceptable. This produces the last row of the data.

An easier way to do this is:

```{r}
hospds %>% slice(41135)
```

or more generally:

```{r}
hospds %>% slice(n())
```

The `n()` function does a count - in this case the number of rows in the data.


### Exercises

1. What is the IQR (interquartile range) of `Age`? You will need the the `IQR()` function.

2. Compute the sum of `Age` and the sum of `TreatDays`. These two numbers can be computed with a single line of code.

3. Extract the last value of all the character type variables. You will need the `is.character()` function.

4. What is the mean of the ten largest values of `TreatDays`?

## Grouping

Summarizing is most useful for groups within the data rather than for all the data. Here is an example:

```{r}
hospds %>% group_by(City) %>% summarise(mean(Age))
```

We group by city and then compute the mean age.


Here is another example:

```{r}
hospds %>% group_by(Age) %>% summarise(meanstay = mean(TreatDays), maxstay = max(TreatDays), n=n())
```

Here we group by age and compute 3 summary statistics: the mean stay, the maximum stay and we use the `n()` function to count the number of patients in each group.

You can group by more than one variable. For example:

```{r}
hospds %>% group_by(City, Occupation) %>% summarize(meanstay = mean(TreatDays), meanage= mean(Age), n=n())
```

We group by occupation and city and compute mean stay and age along with the number in that category.


### Exercises

1. What are the median stays for males and females?

2. What are the median ages for each level of Education, separately by sex?

3. Compute the number of combinations of `Education`, `Occupation` and `City`. Sort them by their frequency.

4. Extract the first of each combination of `Gender`, `Finish`. Your solution should be a dataset with 6 rows.



